# Assignment 1

```python
import numpy as np
```

## 1. (1.0 point) Function return value

This exercise is to show you how function return values work. Run the code below to see what happens.

```python
def f1(x):
    print(x + 1)

def f2(x):
    return(x + 1)

f1(4)
f2(4)
```

Output:
```
5
5
```

Each function should produce the same output. Now run the following code.

```python
f1(4) + 1
f2(4) + 1
```

Output:
```
5
TypeError                                 Traceback (most recent call last)
<ipython-input-4-02362f8b3abb> in <module>()
----> 1 f1(4) + 1
      2 f2(4) + 1

TypeError: unsupported operand type(s) for +: 'NoneType' and 'int'
```

Explain what is happening and why you get this result.

Ans: The type of f1(4) is None, so there will be an error when adding it with an integer.
     But the type of f2(4) is int, so there will be no error.

## 2. (1.0 point) Function operation

Consider the following piece of code:

```python
def f(x):
    if x == 0 or x == 1:
        return x
    return f(x - 1) + f(x - 2)
```

Describe, in words, what this code does, and how it does it.

Ans:
$$f(0) = 0$$
$$f(1) = 1$$     
\begin{align}f(n) = f(n - 1) + f(n - 2),\ for\ all\ integers\ larger\ than\ 1\end{align}

For what inputs will this function not behave as expected? What will happen?

```python
def f(x):
    if x == 0 or x == 1:
        return x
    return f(x - 1) + f(x - 2)
f(-1) # type 1 error: negative input
f(3.4) # type 2 error: the input is not an int
```

Output:
```
... last 1 frames repeated, from the frame below ...

<ipython-input-6-89ca98867d3f> in f(x)
      2     if x == 0 or x == 1:
      3         return x
----> 4     return f(x - 1) + f(x - 2)
      5 f(1.4)

RecursionError: maximum recursion depth exceeded in comparison
```

## 3. (1.0 point) Defining a function

The following program plays the game of "Rock, Paper, and Scissors". Familiarize yourself with how it works. Try running the program.

```python
from random import randint
valid_input = ["R", "P", "S"]
winners = [["R", "S"], ["S", "P"], ["P", "R"]]

print("You are playing a Rock, Paper, Scissors simulation.")
user_choice = input("Select R for rock, P for Paper, S for scissors, Q to quit\n> ")
user_choice = user_choice.upper()

while user_choice != 'Q':
    if user_choice not in valid_input:
        print("Error: invalid input -- " + user_choice)
        break
    
    computer_choice = valid_input[randint(0, len(valid_input)-1)]
    print("Computer chooses " + computer_choice)

    if user_choice == computer_choice:
        print("Tie Game")
    elif [user_choice, computer_choice] in winners:
        print("You win!")
    else:
        print("Computer wins!")
    
    user_choice = input("Select R for rock, P for Paper, S for scissors, Q to quit\n> ")
    user_choice = user_choice.upper()
```

Output:
```
You are playing a Rock, Paper, Scissors simulation.
Select R for rock, P for Paper, S for scissors, Q to quit
> r
Computer chooses S
You win!
Select R for rock, P for Paper, S for scissors, Q to quit
> d
Error: invalid input -- D
```

Recall how we define a function using `def`, and how we pass in parameters. Transform the code into a function called `rps` that takes two parameters: `user_choice` and `computer_choice`, with valid inputs `'R'`, `'P'`, and `'S'`. Your function should return `'D'` for a draw, `'W'` if `user_choice` wins, and `'L'` if `user_choice` loses.

```python
### BEGIN SOLUTION
valid_input = ["R", "P", "S"]
winners = [["R", "S"], ["S", "P"], ["P", "R"]]
def rps(user_choice, computer_choice):
    if user_choice not in valid_input:
        print("Error: invalid input -- " + user_choice)
        return
    if computer_choice not in valid_input:
        print("Error: invalid input -- " + computer_choice)
        return
    if user_choice == computer_choice:
        return 'D'
    elif [user_choice, computer_choice] in winners:
        return 'W'
    else:
        return 'L'
### END SOLUTION
```

<!-- grade_id: Q3C, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(rps('R', 'R'), 'D')
assert_equal(rps('R', 'P'), 'L')
assert_equal(rps('R', 'S'), 'W')
assert_equal(rps('P', 'R'), 'W')
assert_equal(rps('P', 'P'), 'D')
assert_equal(rps('P', 'S'), 'L')
assert_equal(rps('S', 'R'), 'L')
assert_equal(rps('S', 'P'), 'W')
assert_equal(rps('S', 'S'), 'D')
```

## 4. (1.0 point) List function

Given a list of non-empty tuples, return a list sorted in increasing order by the last element in each tuple. 

For example: `[(1, 7), (1, 3), (3, 4, 5), (2, 2)]` yields `[(2, 2), (1, 3), (3, 4, 5), (1, 7)]`

<!-- grade_id: Q4A, points: 0, grade: false, solution: true -->
```python
def sort_tuple_list(tuple_list):
    '''Function to sort a list of non-empty tuples in increasing order by the last element in each tuple'''
    ### BEGIN SOLUTION
    lastele = list()
    result = list()
    for i in range(len(tuple_list)):
        lastele.append(tuple_list[i][-1])
    index = np.argsort(lastele)
    for i in range(len(tuple_list)):
        result.append(tuple_list[index[i]])
    return result
    ### END SOLUTION
```

```python
d = [2, 4, 5, 1]
np.argsort(d)

```

Output:
```
array([3, 0, 1, 2])
```

<!-- grade_id: Q4B, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(sort_tuple_list([(1, 3), (3, 2), (2, 1)]), [(2, 1), (3, 2), (1, 3)])
assert_equal(sort_tuple_list([(2, 3), (1, 2), (3, 1)]), [(3, 1), (1, 2), (2, 3)])
assert_equal(sort_tuple_list([(1, 7), (1, 3), (3, 4, 5), (2, 2)]), [(2, 2), (1, 3), (3, 4, 5), (1, 7)])
```

# 5. (1.0 point) List sort

Given two lists sorted in increasing order, create and return a merged list of all the elements in sorted order. You may modify the passed in lists. Ideally, the solution should work in "linear" time, making a single pass of both lists.

<!-- grade_id: Q5A, points: 0, grade: false, solution: true -->
```python
def list_merge(list1, list2):
    '''Given two lists sorted in increasing order, create and 
       return a merged list of all the elements in sorted order
    '''
    ### BEGIN SOLUTION
    len1 = len(list1)
    len2 = len(list2)
    result = list()
    while (len1 > 0) and (len2 > 0):
        if list1[0] < list2[0]:
            result.append(list1.pop(0))
            len1 = len1 - 1
        else:
            result.append(list2.pop(0))
            len2 = len2 - 1
    if len1 == 0:
        result.extend(list2)
    if len2 == 0:
        result.extend(list1)
    return result
    ### END SOLUTION
```

<!-- grade_id: Q5B, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(list_merge(['aa', 'xx', 'zz'], ['bb', 'cc']), ['aa', 'bb', 'cc', 'xx', 'zz'])
assert_equal(list_merge(['aa', 'xx'], ['bb', 'cc', 'zz']), ['aa', 'bb', 'cc', 'xx', 'zz'])
assert_equal(list_merge(['aa', 'aa'], ['aa', 'bb', 'bb']), ['aa', 'aa', 'aa', 'bb', 'bb'])
```

# 6. (1.0 point) Built-in Operator

Imagine that Python doesn’t have the `!=` operator built in. Write a function that takes two parameters and gives the same result as the `!=` operator. Obviously, you cannot use `!=` within your function! Test if your code works by thinking of examples and making sure the output is the same for your new method as `!=` gives you.

<!-- grade_id: Q6A, points: 0, grade: false, solution: true -->
```python
def not_equal(val1, val2):
    ''' User supplied version of the != operator, implemeted without using !='''
    ### BEGIN SOLUTION
    if val1 == val2:
        return False
    else:
        return True
    ### END SOLUTION
```

<!-- grade_id: Q6B, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_true, assert_false
assert_true(not_equal(2,3))
assert_false(not_equal(3,3))
assert_true(not_equal('3',3))
assert_true(not_equal((1,2),[1,2]))
assert_false(not_equal((3,4),(3,4)))
```

## 7. (1.5 point) Quadratic Formula

A quadratic equation is a second order polynomial with three coefficients a, b, and c. It is given by:

$$ax^2 + bx + c = 0$$

The solutions to the equation are given by the formulas:

$$x_1 = \frac{-b + \sqrt{b^2 - 4ac}}{2a}$$

and

$$x_2 = \frac{-b - \sqrt{b^2 - 4ac}}{2a}$$



Write a function that computes the roots of a quadratic equation and returns them as a tuple. Your function should take three parameters, a, b, and c. Note that if

$$b^2 - 4ac < 0$$ 

then the roots will be complex. Hint: remember that Python has a built-in complex data type.

<!-- grade_id: Q7A, points: 0, grade: false, solution: true -->
```python
def roots(a, b, c):
    ''' Computes the roots of the quadratic equation with coefficients a, b, c, and returns
        them as the tuple (x1, x2)
    '''
    ### BEGIN SOLUTION
    res = tuple()
    if a == 0 and b != 0:
        res = res +(-c/b,-c/b)
    if a == 0 and b == 0:
        res = res + (complex(0,0), complex(0,0))
    if a != 0:
        deltasqr = b ** 2 - 4 * a * c
        if deltasqr >= 0:
            res = res + ((-b + np.sqrt(deltasqr)) / (2 * a), (-b - np.sqrt(deltasqr)) / (2 * a))
        else:
            res = res + (complex(-b/(2*a), np.sqrt(-deltasqr)/(2 * a)), complex(-b/(2*a), -np.sqrt(-deltasqr)/(2 * a)))
    return res
    ### END SOLUTION
```

<!-- grade_id: Q7B, points: 2, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(roots(1,2,3), ((-1+1.4142135623730951j), (-1-1.4142135623730951j)))
assert_equal(roots(0,0,4), (0j, 0j))
assert_equal(roots(-4,3,2), (-0.42539052967910607, 1.175390529679106))
```

## 8. (2.5 points) The Game of Nims

In this game, two players sit in front of a pile of 100 stones. They take turns, each removing between 1 and 5 stones (assuming there are at least 5 stones left in the pile). The person who removes the last stone(s) wins.

In this problem, you’ll write a function to play this game. It may seem tricky, so break it down into parts. Like many programs, we have to use nested loops (one loop inside another). In the outermost loop, we want to keep playing until we are out of stones. Inside that, we want to keep alternating players. You have the option of either writing two blocks of code, or keeping a variable that tracks the current player.

We are going to use the `input` function to obtain the user's move. We might want to have an innermost loop that checks if the user’s input is valid. Is it a number? Is it a valid number (e.g. between 1 and 5)? Are there enough stones in the pile to take off this many? If any of these answers are no, we should tell the user and re-ask them the question.

If you choose to write two blocks of code, the basic outline of the program should be something like this:

    while [pile is not empty]:
       while [player 1’s answer is not valid]:
          [ask player 1]
       [execute player 1’s move]
       [same as above for player 2]
       
Be careful with the validity checks. Specifically, we want to keep asking player 1 for their choice as long as their answer is not valid, BUT we want to make sure we ask them at least ONCE. So, for example, we will want to keep a variable that tracks whether their answer is valid, and set it to `False` initially.

There is no need to limit yourself to only one function! Use as many functions as you want in order to make the code more readable and modular.

When you're finished, test each other's programs by playing them!

<!-- grade_id: Q8A, points: 4, grade: true, solution: true -->
```python

def play_nims(pile, max_stones):
    '''
    An interactive two-person game; also known as Stones.
    @param pile: the number of stones in the pile to start
    @param max_stones: the maximum number of stones you can take in one turn
    '''

## Basic structure of program (feel free to alter as you please):

#    while [pile is not empty]:
#        while [player 1's answer is not valid]:
#            [ask player 1]
#            [execute player 1's move]
#       
#        while [player 2's answer is not valid]:
#            [ask player 2]
#            [execute player 2's move]
#
#    print "Game over"
### BEGIN SOLUTION
    endflag = False
    while endflag != True:
        p1 = ask(pile, max_stones)
        pile = pile - p1
        endflag = (pile <= 0)
        if endflag == True:
            print('player 1 wins!')
            break
        p2 = ask(pile, max_stones)
        pile = pile - p2
        endflag = (pile <= 0)
        if endflag == True:
            print('player 2 wins!')
            break
        
            
        
### END SOLUTION
```

```python
def ask(pile, max_stones):
    isvalid = False
    while isvalid != True:
        currmax = min(pile, max_stones)
        validans = list(range(currmax + 1))[1:]
        print('valid inputs:')
        print(validans)
        print('current total number of rocks:', pile)
        inp1 = input("Select the number of rocks you want to take\n> ")
        if('.' in inp1):
            print('error, it should be an integer!')
            continue
        if int(inp1) in validans:
            isvalid = True
            return int(inp1)

      
```

```python
play_nims(30, 5)
```

Output:
```
valid inputs:
[1, 2, 3, 4, 5]
current total number of rocks: 30
Select the number of rocks you want to take
> 10
valid inputs:
[1, 2, 3, 4, 5]
current total number of rocks: 30
Select the number of rocks you want to take
> 4
valid inputs:
[1, 2, 3, 4, 5]
current total number of rocks: 26
Select the number of rocks you want to take
> 3
valid inputs:
[1, 2, 3, 4, 5]
current total number of rocks: 23
Select the number of rocks you want to take
> 5
valid inputs:
[1, 2, 3, 4, 5]
current total number of rocks: 18
Select the number of rocks you want to take
> 4
valid inputs:
[1, 2, 3, 4, 5]
current total number of rocks: 14
Select the number of rocks you want to take
> 4
valid inputs:
[1, 2, 3, 4, 5]
current total number of rocks: 10
Select the number of rocks you want to take
> 3
valid inputs:
[1, 2, 3, 4, 5]
current total number of rocks: 7
Select the number of rocks you want to take
> 4
valid inputs:
[1, 2, 3]
current total number of rocks: 3
Select the number of rocks you want to take
> 4
```

```python

```